---
title: 'Chapter Title Here'
description: 'Chapter description goes here.'
free_preview: true
---

## Example coding exercise

```yaml
type: NormalExercise
key: e8c1edbe67
lang: python
xp: 100
skills: 2
```

This is an example exercise.

`@instructions`


`@hint`


`@pre_exercise_code`
```{python}

```

`@sample_code`
```{python}

```

`@solution`
```{python}

```

`@sct`
```{python}

```

---

## Evaluating Logistic Regression

```yaml
type: VideoExercise
key: 6bb7f7a35b
xp: 50
```

`@projector_key`
2c019d1e7d05eda0ad12c0130582d6ce
